# Name
Please write your name down here:

<!-- grade_id: cell-758d6c729dc4c726, points: 0, grade: true, solution: true -->
Peter Brederlow

**Pandas Basics**

In this notebook we will familiarize ourselves with some basic `pandas` objects and operations which we'll use often later in the practical course.

Pandas has extensive online documentation, and you may find their [10 Minutes Introduction](http://pandas.pydata.org/pandas-docs/stable/10min.html) and their [tutorials](http://pandas.pydata.org/pandas-docs/stable/tutorials.html) especially useful.

Also take a look at [Jupyter keyboard shortcuts](https://www.cheatography.com/weidadeyue/cheat-sheets/jupyter-notebook/) to help you format your notebooks efficiently.

```python
import pandas as pd
import numpy as np
pd.set_option('future.no_silent_downcasting', True)
```

Output:
```
/var/folders/vm/8qbbn1f91lb9f4yl4gd82lch0000gn/T/ipykernel_98364/1153847508.py:3: Pandas4Warning: 'future.no_silent_downcasting' is deprecated, please refrain from using it.
  pd.set_option('future.no_silent_downcasting', True)
```

<!-- grade_id: cell-90e9c5984ef9c1b7, points: 0, grade: true, solution: false -->
```python
# Test declarations, do not delete!
```

# Series objects

## Create a pandas `Series` object named `s` with the following data:

| s   |
|-----|
| 3.0 |
| 1.5 |
| 4.3 |
| 8.2 |
| 0.9 |

<!-- grade_id: cell-5d304cbec278f0d6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
s = pd.Series([3.0, 1.5, 4.3, 8.2, 0.9])
```

<!-- grade_id: cell-57c63531adae75e2, points: 0.25, grade: true, solution: false -->
```python
assert 's' in locals()
# test for type
print(s.dtype)
```

Output:
```
float64
```

<!-- grade_id: cell-75074cb0fab960b8, points: 0.25, grade: true, solution: false -->
```python
# test for content
print(s)
```

Output:
```
0    3.0
1    1.5
2    4.3
3    8.2
4    0.9
dtype: float64
```

## Access its first value
And store it as variable `first`.

<!-- grade_id: cell-75db399d877ca9ac, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
first = s[0]
print(first)
```

Output:
```
3.0
```

<!-- grade_id: cell-838176581d571899, points: 0.25, grade: true, solution: false -->
```python
assert 'first' in locals()
# test for value
print(first.dtype)
```

Output:
```
float64
```

## Access its last value
And store it as variable `last`.

<!-- grade_id: cell-d49ec356466020a0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
last = s.iloc[-1]
print(last)
```

Output:
```
0.9
```

<!-- grade_id: cell-2c7634e8d4b7bb57, points: 0.25, grade: true, solution: false -->
```python
assert 'last' in locals()
# test for value
print(last.dtype)
```

Output:
```
float64
```

## Sort it in ascending and descending order
Store the ascending series as `asc` and the descending series as `desc`.

<!-- grade_id: cell-d266bcc1a956b0eb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
asc = s.sort_values()
print(asc)

desc = s.sort_values(ascending=False)
print(desc)
```

Output:
```
4    0.9
1    1.5
0    3.0
2    4.3
3    8.2
dtype: float64
3    8.2
2    4.3
0    3.0
1    1.5
4    0.9
dtype: float64
```

<!-- grade_id: cell-bae1080de844d85f, points: 0.25, grade: true, solution: false -->
```python
assert 'asc' in locals()
assert type(asc) == pd.Series
```

<!-- grade_id: cell-b8b72d9fb834c51d, points: 0.25, grade: true, solution: false -->
```python
assert 'desc' in locals()
assert type(desc) == pd.Series
```

## Calculate its mean and standard deviation
Store the mean in `avg` and the standard deviation in `stdev`.

<!-- grade_id: cell-0c088f0359b747e0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
avg = s.mean()
print(avg)
stdev = s.std()
print(stdev)
```

Output:
```
3.5799999999999996
2.90292955477738
```

<!-- grade_id: cell-8ea4b1d2de1f0482, points: 0.25, grade: true, solution: false -->
```python
assert 'avg' in locals()
```

<!-- grade_id: cell-4496982455bf36ea, points: 0.25, grade: true, solution: false -->
```python
assert 'stdev' in locals()
```

# DataFrame objects, accessing values

Create a `DataFrame` named `df_species` of the following table:

|(index) |species |weight  |legs |
|--------|--|---|----|
|0       |human |62  |2 |
|1       |cat |4  |4 |
|2       |mouse |0.02  |4 |
|3       |dog |20 |4 |
|4       |mole |0.09  |4 |
|5       |train |200000 |NaN |
|6       |bee |0.0001 |6  |
|7       |elephant |3000 |4 |

<div class="alert alert-block alert-warning"><b>Warning: </b> Do not use <code>pd.NA</code> to represent missing values, because this will break the behavior of multiple functions. </div>

<!-- grade_id: cell-6a32c0410086bf9b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df_species = pd.DataFrame({'species': ['human', 'cat', 'mouse', 'dog', 'mole', 'train', 'bee', 'elephant'], 
                           'weight': [62, 4, 0.02, 20, 0.09, 200000, 0.0001, 3000],
                           'legs': [2, 4, 4, 4, 4, np.nan, 6, 4]})
print(df_species)
```

Output:
```
species       weight  legs
0     human      62.0000   2.0
1       cat       4.0000   4.0
2     mouse       0.0200   4.0
3       dog      20.0000   4.0
4      mole       0.0900   4.0
5     train  200000.0000   NaN
6       bee       0.0001   6.0
7  elephant    3000.0000   4.0
```

<!-- grade_id: cell-ad11ed68dd452a93, points: 1, grade: true, solution: false -->
```python
assert 'df_species' in locals()
# test for content
```

## Sort `df_species` by `weight` and overwrite it.

<!-- grade_id: cell-47664d9b0619f9c5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df_species = df_species.sort_values('weight') #assuming ascending order
```

<!-- grade_id: cell-bc5a534439a283ef, points: 0.5, grade: true, solution: false -->
```python
assert 'df_species' in locals()
# test for content
```

## Access the 3rd row of `df_species` as a `Series` object
And store it in `third`.

<!-- grade_id: cell-74b4d6a5899a575b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
third = df_species.iloc[2]
print(third)
```

Output:
```
species             mole
weight              0.09
legs                 4.0
weight_in_lb    0.198416
Name: 4, dtype: object
```

<!-- grade_id: cell-75f5a029c585d9db, points: 0.5, grade: true, solution: false -->
```python
assert 'third' in locals()
assert type(third) == pd.Series
```

## Access the row with INDEX `3` of `df_species`
And store it in `third_index`.

<!-- grade_id: cell-3b4fbf3b595f863a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
third_index = df_species.loc[3]
print(third_index)
```

Output:
```
species     dog
weight     20.0
legs        4.0
Name: 3, dtype: object
```

<!-- grade_id: cell-44e8dfaca541ca59, points: 0.5, grade: true, solution: false -->
```python
assert 'third_index' in locals()
assert type(third_index) == pd.Series
```

## Access the column `weight`
And store it in `weight`.

<!-- grade_id: cell-e62fd18d26ed5ba7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
weight = df_species['weight']
print(weight)
```

Output:
```
6         0.0001
2         0.0200
4         0.0900
1         4.0000
3        20.0000
0        62.0000
7      3000.0000
5    200000.0000
Name: weight, dtype: float64
```

<!-- grade_id: cell-b9fd9b3c84d68b0e, points: 0.25, grade: true, solution: false -->
```python
assert 'weight' in locals()
assert type(weight) == pd.Series
```

## Access the index and column labels
And store them as `index` and `columns`.

<!-- grade_id: cell-5155d4f698df44f5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
index = df_species.index
columns = df_species.columns
print(index)
print(columns)
```

Output:
```
Index([6, 2, 4, 1, 3, 0, 7, 5], dtype='int64')
Index(['species', 'weight', 'legs'], dtype='str')
```

<!-- grade_id: cell-88a2969c91f8e3c0, points: 0.25, grade: true, solution: false -->
```python
assert 'index' in locals()
assert type(index) == pd.Index
```

<!-- grade_id: cell-40d81a0027df5f27, points: 0.25, grade: true, solution: false -->
```python
assert 'columns' in locals()
assert type(columns) == pd.Index
```

## Extract the third row up until the fifth row into `df_slice`

<!-- grade_id: cell-6ecb8a3bc3b8724a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df_slice = df_species.iloc[2:4]
print(df_slice)
```

Output:
```
species  weight  legs
4    mole    0.09   4.0
1     cat    4.00   4.0
```

<!-- grade_id: cell-8073bb83d386662a, points: 0.25, grade: true, solution: false -->
```python
assert 'df_slice' in locals()
assert type(df_slice) == pd.DataFrame
```

## Extract rows without missing values into a new DataFrame into `df_non_empty`

<!-- grade_id: cell-c240022128111dbd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df_non_empty = df_species.dropna()
df_non_empty
```

Output:
```
species     weight  legs
6       bee     0.0001   6.0
2     mouse     0.0200   4.0
4      mole     0.0900   4.0
1       cat     4.0000   4.0
3       dog    20.0000   4.0
0     human    62.0000   2.0
7  elephant  3000.0000   4.0
```

<!-- grade_id: cell-2952c1f81bee5471, points: 0.25, grade: true, solution: false -->
```python
assert 'df_non_empty' in locals()
```

## Obtain a summary table of the DataFrame with the means, standard deviations, min/max/percentiles of the numerical columns
Store it in `df_summary`

<div class="alert alert-block alert-info"><b>Hint: pandas has a nice convencience function for this.</</a></div>

<!-- grade_id: cell-d94c2b0ab7b40775, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df_summary = df_species.describe()
df_summary
```

Output:
```
weight      legs
count       8.000000  7.000000
mean    25385.763762  4.000000
std     70562.543935  1.154701
min         0.000100  2.000000
25%         0.072500  4.000000
50%        12.000000  4.000000
75%       796.500000  4.000000
max    200000.000000  6.000000
```

<!-- grade_id: cell-ecdb6c423e0ff29a, points: 0.5, grade: true, solution: false -->
```python
assert 'df_summary' in locals()
```

## Add a new column `weight_in_lb` to the DataFrame containing the weights in pounds instead of kilograms
**For the sake of simplicity limit yourself to the first decimal of pounds when converting the weight.**

<!-- grade_id: cell-c1c7ad03cf14a3f2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df_species['weight_in_lb'] = df_species['weight'] * 2.20462
df_species
```

Output:
```
species       weight  legs   weight_in_lb
6       bee       0.0001   6.0       0.000220
2     mouse       0.0200   4.0       0.044092
4      mole       0.0900   4.0       0.198416
1       cat       4.0000   4.0       8.818480
3       dog      20.0000   4.0      44.092400
0     human      62.0000   2.0     136.686440
7  elephant    3000.0000   4.0    6613.860000
5     train  200000.0000   NaN  440924.000000
```

<!-- grade_id: cell-3410473b16e320a4, points: 0.5, grade: true, solution: false -->
```python
assert 'weight_in_lb' in df_species.columns
```

# Data Import

## Import the `df_characters.csv` file into a DataFrame `df_characters`, and display it.
Use the url from GitHub directly: `https://raw.githubusercontent.com/Practical-Integrative-Bioinformatics/Introduction/refs/heads/main/example_data/df_characters.csv`.

<!-- grade_id: cell-a9c1a31a13761fdd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
url = 'https://raw.githubusercontent.com/Practical-Integrative-Bioinformatics/Introduction/refs/heads/main/example_data/df_characters.csv'
df_characters = pd.read_csv(url)
df_characters
```

Output:
```
name country   species       weight  birthyear
0               Pumuckl     GER    kobold       2.0000       1982
1                   Tom     USA       cat       3.0000       1940
2                 Jerry     USA     mouse       0.3000       1940
3          Mickey Mouse     USA     mouse      10.0000       1928
4               Asterix     FRA     human      58.0000       1959
5                Obelix     FRA     human     180.0000       1959
6           Hello Kitty     JPN       cat       0.7000       1974
7               Songoku     JPN     human      62.0000       1984
8        Maulwurf Krtek     CZE      mole       0.1000       1956
9                Tintin     BEL     human      54.0000       1929
10               Thomas     GBR     train  115000.0000       1984
11           Biene Maja     GER       bee       0.0001       1975
12                 Gary     USA     snail       0.0500       1999
13                Babar     CAN  elephant    2000.0000       1989
14    Benjamin Blümchen     G
```

<!-- grade_id: cell-40a964999f209bb3, points: 0.5, grade: true, solution: false -->
```python
assert 'df_characters' in locals()
```

## Now import the `df_countries.tsv` into a pandas DataFrame as `df_countries`. Do you need to change any parameters of the importing function?

Use the url from GitHub directly: `https://raw.githubusercontent.com/Practical-Integrative-Bioinformatics/Introduction/refs/heads/main/example_data/df_countries.tsv`.

<!-- grade_id: cell-d66313c822446ac5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
url = 'https://raw.githubusercontent.com/Practical-Integrative-Bioinformatics/Introduction/refs/heads/main/example_data/df_countries.tsv'
df_countries = pd.read_csv(url, sep='\t')
df_countries
```

Output:
```
code                   country continent  population
0  GER                   Germany    Europe        83.0
1  USA  United States of America   America       327.0
2  FRA                    France    Europe        67.0
3  JPN                     Japan      Asia       127.0
4  CZE            Czech Republic    Europe        10.5
5  BEL                   Belgium    Europe        11.3
6  GBR            United Kingdom    Europe        66.0
7  CHN                     China      Asia      1386.0
8  CAN                    Canada   America        37.0
9  BRA                    Brazil   America       210.0
```

<!-- grade_id: cell-553a2db9d29a6e84, points: 0.5, grade: true, solution: false -->
```python
assert 'df_countries' in locals()
```

# DataFrame filtering
## Using `df_characters` apply the following filters (separately) and display the results:

### All German characters
Store the result in `Germans`.

<!-- grade_id: cell-9a0257acd9ca8cb8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
germans = df_characters[df_characters['country'] == 'GER']
print(germans)
```

Output:
```
name country   species    weight  birthyear
0               Pumuckl     GER    kobold    2.0000       1982
11           Biene Maja     GER       bee    0.0001       1975
14    Benjamin Blümchen     GER  elephant  800.0000       1988
15  Prinzessin Lillifee     GER     human   30.0000       2004
```

<!-- grade_id: cell-ae0a5b45cf413692, points: 0.5, grade: true, solution: false -->
```python
assert 'germans' in locals()
```

### Characters born after 1970
Store the result in `seventies`.

<!-- grade_id: cell-3fe81aa4f90c1545, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
seventies = df_characters[df_characters['birthyear'] >= 1970]
print(seventies)
```

Output:
```
name country   species       weight  birthyear
0               Pumuckl     GER    kobold       2.0000       1982
6           Hello Kitty     JPN       cat       0.7000       1974
7               Songoku     JPN     human      62.0000       1984
10               Thomas     GBR     train  115000.0000       1984
11           Biene Maja     GER       bee       0.0001       1975
12                 Gary     USA     snail       0.0500       1999
13                Babar     CAN  elephant    2000.0000       1989
14    Benjamin Blümchen     GER  elephant     800.0000       1988
15  Prinzessin Lillifee     GER     human      30.0000       2004
```

<!-- grade_id: cell-c02739f4a1e5448c, points: 0.5, grade: true, solution: false -->
```python
assert 'seventies' in locals()
```

### Human characters weighing <60 kg
Store the result in `less_than_sixty`.

<!-- grade_id: cell-9369d709d847b616, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
less_than_sixty = df_characters[df_characters['weight'] < 60]
print(less_than_sixty)
```

Output:
```
name country species   weight  birthyear
0               Pumuckl     GER  kobold   2.0000       1982
1                   Tom     USA     cat   3.0000       1940
2                 Jerry     USA   mouse   0.3000       1940
3          Mickey Mouse     USA   mouse  10.0000       1928
4               Asterix     FRA   human  58.0000       1959
6           Hello Kitty     JPN     cat   0.7000       1974
8        Maulwurf Krtek     CZE    mole   0.1000       1956
9                Tintin     BEL   human  54.0000       1929
11           Biene Maja     GER     bee   0.0001       1975
12                 Gary     USA   snail   0.0500       1999
15  Prinzessin Lillifee     GER   human  30.0000       2004
```

<!-- grade_id: cell-0b91adb38f2e9892, points: 0.5, grade: true, solution: false -->
```python
assert 'less_than_sixty' in locals()
```

## Introduce a column named `age` based on the characters' birthdates.

<!-- grade_id: cell-13e68e047462cb6d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df_characters['age'] = 2026 - df_characters['birthyear']
```

<!-- grade_id: cell-a3361108f1c507ed, points: 0.5, grade: true, solution: false -->
```python
assert 'age' in df_characters.columns
```

## What is the average age of the characters?
Store it in `avg_age`.

<!-- grade_id: cell-b51a19b878aa9892, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
avg_age = df_characters['age'].mean()
print(avg_age)
```

Output:
```
57.875
```

<!-- grade_id: cell-10dfb3933902e500, points: 0.5, grade: true, solution: false -->
```python
assert 'avg_age' in locals()
```

## Who is the youngest German character?
Store the corresponding row in `youngest_german`.

<!-- grade_id: cell-27be1f6024d8505a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
germans = df_characters[df_characters['country'] == 'GER']
youngest_german = germans.loc[germans['age'].idxmin()]
print(youngest_german)
```

Output:
```
name         Prinzessin Lillifee
country                      GER
species                    human
weight                      30.0
birthyear                   2004
age                           22
Name: 15, dtype: object
```

<!-- grade_id: cell-6dd5fd5983a024cd, points: 0.5, grade: true, solution: false -->
```python
assert 'youngest_german' in locals()
```

# DataFrame aggregation

Using `df_characters`

## Count the number of characters from each country
And store it as series `chars_per_country`

<!-- grade_id: cell-8c146f7eb7b0b706, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
chars_per_country = df_characters['country'].value_counts()
print(chars_per_country)
```

Output:
```
country
GER    4
USA    4
FRA    2
JPN    2
CZE    1
BEL    1
GBR    1
CAN    1
Name: count, dtype: int64
```

<!-- grade_id: cell-1fb7917c96cec2da, points: 0.5, grade: true, solution: false -->
```python
assert 'chars_per_country' in locals()
```

## Calculate the average weight and age of characters by species
Store the result in `df_characters_avg`. The new dataframe should only contain the columns `species`, `avg_weight` and `avg_age`.

<!-- grade_id: cell-09ccda92aaeadbdd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df_characters_avg = df_characters.groupby('species', as_index=False).agg(
    avg_weight=('weight', 'mean'),
    avg_age=('age', 'mean'))
df_characters_avg
```

Output:
```
species   avg_weight  avg_age
0       bee       0.0001     51.0
1       cat       1.8500     69.0
2  elephant    1400.0000     37.5
3     human      76.8000     59.0
4    kobold       2.0000     44.0
5      mole       0.1000     70.0
6     mouse       5.1500     92.0
7     snail       0.0500     27.0
8     train  115000.0000     42.0
```

<!-- grade_id: cell-ddac6f79aedb1e51, points: 0.5, grade: true, solution: false -->
```python
assert 'df_characters_avg' in locals()
assert 'avg_weight' in df_characters_avg.columns
assert 'avg_age' in df_characters_avg.columns
```

## List the heaviest characters per country
Store the result in `df_chars_heaviest`. The new dataframe should only contain the columns `country`, `name`, `weight`.

<!-- grade_id: cell-3e607e3ad3add1eb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
heaviest_indices = df_characters.groupby('country')['weight'].idxmax()
df_chars_heaviest = df_characters.loc[heaviest_indices, ['country', 'name', 'weight']]
df_chars_heaviest
```

Output:
```
country               name    weight
9      BEL             Tintin      54.0
13     CAN              Babar    2000.0
8      CZE     Maulwurf Krtek       0.1
5      FRA             Obelix     180.0
10     GBR             Thomas  115000.0
14     GER  Benjamin Blümchen     800.0
7      JPN            Songoku      62.0
3      USA       Mickey Mouse      10.0
```

<!-- grade_id: cell-30be9b2f624dc5d4, points: 0.5, grade: true, solution: false -->
```python
assert 'df_chars_heaviest' in locals()
assert 'name' in df_chars_heaviest.columns
assert 'weight' in df_chars_heaviest.columns
```

# Joining (or merging) DataFrames
We now have three dataframes: `df_species`, `df_characters` and `df_countries`. Take another look at all of them and try to come up with a way to combine their information.
Create another dataframe called `df_merge` including all three dataframes. There are some things you should consider, when you create your new dataframe:
1. We are going to use our characters from `df_characters` as main target, so leave out all countries or species without one.
2. Character and species each have weights, use `weight_indiv` for a characters weight and `weight_species` for the species weight.
3. Countries have a column `country` and `code`, think about which you should use as a key to join your dataframes.
4. Try to solve this task using (mostly) the `merge` function in pandas.

<div class="alert alert-block alert-info"><b>Hint: Should you struggle or just want to compare your result, you can take a look at the merged dataframe <a href=https://github.com/Practical-Integrative-Bioinformatics/Introduction/blob/main/example_data/day1.tsv>here</a>.</</a></div> Watch out, as this is not going to look exactly like your dataframe, as this contains additional columns and different ages.

<!-- grade_id: cell-bb8918cd68bcd153, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
#inner merge, since only countries and species that are in df_characters are kept
df_merge = pd.merge(df_characters, df_countries, left_on='country', right_on='code', how = 'left')
df_merge
df_merge = pd.merge(df_merge, df_species, left_on='species', right_on='species', how = 'left')
df_merge
```

Output:
```
name country_x   species     weight_x  birthyear  age code  \
0               Pumuckl       GER    kobold       2.0000       1982   44  GER   
1                   Tom       USA       cat       3.0000       1940   86  USA   
2                 Jerry       USA     mouse       0.3000       1940   86  USA   
3          Mickey Mouse       USA     mouse      10.0000       1928   98  USA   
4               Asterix       FRA     human      58.0000       1959   67  FRA   
5                Obelix       FRA     human     180.0000       1959   67  FRA   
6           Hello Kitty       JPN       cat       0.7000       1974   52  JPN   
7               Songoku       JPN     human      62.0000       1984   42  JPN   
8        Maulwurf Krtek       CZE      mole       0.1000       1956   70  CZE   
9                Tintin       BEL     human      54.0000       1929   97  BEL   
10               Thomas       GBR     train  115000.0000       1984   42  GBR   
11           Biene Maja       GER       bee    
```

<!-- grade_id: cell-d517e3ff4ab30ed3, points: 1, grade: true, solution: false -->
```python
assert 'df_merge' in locals()
```

## Using our new `df_merge` try to answer: what is the average age of four-legged characters?
Store the result in `avg_age_four_legs`.

<!-- grade_id: cell-03b0b99f7ca9dbd0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
avg_age_four_legs = df_merge[df_merge['legs'] == 4]['age'].mean()
avg_age_four_legs
```

Output:
```
np.float64(66.71428571428571)
```

<!-- grade_id: cell-cbf373dc536ec589, points: 0.5, grade: true, solution: false -->
```python
assert 'avg_age_four_legs' in locals()
```

## What is the average weight of four-legged characters?
Store the result in `avg_weight_four_legs`.

<!-- grade_id: cell-a38a12e1cb84b996, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
avg_weight_four_legs = df_merge[df_merge['legs'] == 4]['weight_x'].mean()
avg_weight_four_legs

```

Output:
```
np.float64(402.0142857142857)
```

<!-- grade_id: cell-3e2a6ea96b684a3e, points: 0.5, grade: true, solution: false -->
```python
assert 'avg_weight_four_legs' in locals()
```

## Compute the characters' relative weight as a percentage of the typical weight of their species and add it as a new column `rel_weight`. Are there any characters you couldn't compute this for?

<!-- grade_id: cell-3a95b08f42ea8518, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df_merge['rel_weight'] = df_merge['weight_x'] / df_merge['weight_y']
df_merge
```

Output:
```
name country_x   species     weight_x  birthyear  age code  \
0               Pumuckl       GER    kobold       2.0000       1982   44  GER   
1                   Tom       USA       cat       3.0000       1940   86  USA   
2                 Jerry       USA     mouse       0.3000       1940   86  USA   
3          Mickey Mouse       USA     mouse      10.0000       1928   98  USA   
4               Asterix       FRA     human      58.0000       1959   67  FRA   
5                Obelix       FRA     human     180.0000       1959   67  FRA   
6           Hello Kitty       JPN       cat       0.7000       1974   52  JPN   
7               Songoku       JPN     human      62.0000       1984   42  JPN   
8        Maulwurf Krtek       CZE      mole       0.1000       1956   70  CZE   
9                Tintin       BEL     human      54.0000       1929   97  BEL   
10               Thomas       GBR     train  115000.0000       1984   42  GBR   
11           Biene Maja       GER       bee    
```

<!-- grade_id: cell-8de816922a9a4912, points: 0.5, grade: true, solution: false -->
```python
assert 'rel_weight' in df_merge.columns
```

## How many characters are there from each continent?
Store the result in `chars_per_continent`.

<!-- grade_id: cell-bc7a8056135c26f5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
chars_per_continent = df_merge['continent'].value_counts()
chars_per_continent
```

Output:
```
continent
Europe     9
America    5
Asia       2
Name: count, dtype: int64
```

<!-- grade_id: cell-4fb87470cd8b1c66, points: 0.5, grade: true, solution: false -->
```python
assert 'chars_per_continent' in locals()
```

## Which character comes from the smallest country (by population)?
Store the row in `smallest_country`.

```python
# character from the smallest country
#dfm.sort_values('population').iloc[0]
#dfm.loc[dfm['population'] == dfm['population'].min()]
smallest_country = df_merge.loc[df_merge['population'].idxmin()]
smallest_country
```

Output:
```
name            Maulwurf Krtek
country_x                  CZE
species                   mole
weight_x                   0.1
birthyear                 1956
age                         70
code                       CZE
country_y       Czech Republic
continent               Europe
population                10.5
weight_y                  0.09
legs                       4.0
weight_in_lb          0.198416
rel_weight            1.111111
Name: 8, dtype: object
```

<!-- grade_id: cell-55ee64f07f83becb, points: 0.5, grade: true, solution: false -->
```python
assert 'smallest_country' in locals()
```

## Who is the most overweight European character?
Store the row in `overweight_european`.

<!-- grade_id: cell-e39d80661792e3b0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
overweight_european = df_merge.loc[df_merge.loc[df_merge['continent'] == 'Europe', 'rel_weight'
].idxmax()]
overweight_european
```

Output:
```
name               Obelix
country_x             FRA
species             human
weight_x            180.0
birthyear            1959
age                    67
code                  FRA
country_y          France
continent          Europe
population           67.0
weight_y             62.0
legs                  2.0
weight_in_lb    136.68644
rel_weight       2.903226
Name: 5, dtype: object
```

<!-- grade_id: cell-372c2a4dafc38714, points: 0.5, grade: true, solution: false -->
```python
assert 'overweight_european' in locals()
```

# Data export

## Create a full outer merge of the three DataFrames with question marks in place of N/A values
Store it as `df_export`.

<!-- grade_id: cell-1214a950add155fe, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df_export = pd.merge(df_characters, df_countries, left_on='country', right_on='code', how = 'outer')
df_export = pd.merge(df_export, df_species, left_on='species', right_on='species', how = 'outer')
df_export = df_export.fillna('?')
df_export
```

Output:
```
name country_x   species  weight_x birthyear   age code  \
0            Biene Maja       GER       bee    0.0001    1975.0  51.0  GER   
1           Hello Kitty       JPN       cat       0.7    1974.0  52.0  JPN   
2                   Tom       USA       cat       3.0    1940.0  86.0  USA   
3                     ?         ?       dog         ?         ?     ?    ?   
4                 Babar       CAN  elephant    2000.0    1989.0  37.0  CAN   
5     Benjamin Blümchen       GER  elephant     800.0    1988.0  38.0  GER   
6                Tintin       BEL     human      54.0    1929.0  97.0  BEL   
7               Asterix       FRA     human      58.0    1959.0  67.0  FRA   
8                Obelix       FRA     human     180.0    1959.0  67.0  FRA   
9   Prinzessin Lillifee       GER     human      30.0    2004.0  22.0  GER   
10              Songoku       JPN     human      62.0    1984.0  42.0  JPN   
11              Pumuckl       GER    kobold       2.0    1982.0  44.0  GER   
12   
```

<!-- grade_id: cell-c39d104934818ec1, points: 0.5, grade: true, solution: false -->
```python
assert 'df_export' in locals()
```

# Optional extra tasks:

## List characters that are the sole entries of their country
Store the result in `sole_chars`.

<!-- grade_id: cell-6e469106dd01f717, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
country_counts = df_merge['country_x'].value_counts()
sole_countries = country_counts[country_counts == 1].index
sole_chars = df_merge[df_merge['country_x'].isin(sole_countries)]
sole_chars
```

Output:
```
name country_x   species  weight_x  birthyear  age code  \
8   Maulwurf Krtek       CZE      mole       0.1       1956   70  CZE   
9           Tintin       BEL     human      54.0       1929   97  BEL   
10          Thomas       GBR     train  115000.0       1984   42  GBR   
13           Babar       CAN  elephant    2000.0       1989   37  CAN   

         country_y continent  population   weight_y  legs   weight_in_lb  \
8   Czech Republic    Europe        10.5       0.09   4.0       0.198416   
9          Belgium    Europe        11.3      62.00   2.0     136.686440   
10  United Kingdom    Europe        66.0  200000.00   NaN  440924.000000   
13          Canada   America        37.0    3000.00   4.0    6613.860000   

    rel_weight  
8     1.111111  
9     0.870968  
10    0.575000  
13    0.666667
```

<!-- grade_id: cell-37226d5cf6ca1b60, points: 0, grade: true, solution: false -->
```python
assert 'sole_chars' in locals()
```

## Count characters by decade of birth -- without explicitly inserting a new column
Store the resulting series in `decades`.

<!-- grade_id: cell-ecacb7557fd5ab5d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
decades = (df_characters['birthyear'] // 10 * 10).value_counts().sort_index()
decades
```

Output:
```
birthyear
1920    2
1940    2
1950    3
1970    2
1980    5
1990    1
2000    1
Name: count, dtype: int64
```

<!-- grade_id: cell-66a26a7e62cccae8, points: 0, grade: true, solution: false -->
```python
assert 'decades' in locals()
```